Visual experience shapes functional connectivity between occipital and non-visual networks

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

The present study evaluates the role of visual experience in shaping functional correlations between human extrastriate visual cortex and frontal regions. The authors used fMRI to assess "resting-state" temporal correlations in three groups: sighted adults, congenitally blind adults, and neonates. Previous research has already demonstrated differences in functional correlations between visual and frontal regions in sighted compared to early blind individuals. The novel contribution of the current study lies in the inclusion of an infant dataset, which allows for an assessment of the developmental origins of these differences.

The main results of the study reveal that correlations between prefrontal and visual regions are more prominent in the blind and infant groups, with the blind group exhibiting greater lateralization. Conversely, correlations between visual and somato-motor cortices are more prominent in sighted adults. Based on these data, the authors conclude that visual experience plays an instructive role in shaping these cortical networks. This study provides valuable insights into the impact of visual experience on the development of functional connectivity in the brain.

Strengths:

The dissociations in functional correlations observed among the sighted adult, congenitally blind, and neonate groups provide strong support for the main conclusion regarding postnatal experience-driven shaping of visual-frontal connectivity.

The inclusion of neonates offers a unique and valuable developmental anchor for interpreting divergence between blind and sighted adults. This is a major advance over prior studies limited to adult comparisons.

Convergence with prior findings in the blind and sighted adult groups reinforces the reliability and external validity of the present results.

The split-half reliability analysis in the infant data increases confidence in the robustness of the reported group differences.

Weaknesses:

The manuscript risks overstating a mechanistic distinction between sighted and blind development by framing visual experience as "instructive" and blindness as "reorganizing." Similarly, the binary framing of visual experience and blindness as independent may oversimplify shared plasticity mechanisms.

The interpretation of changes in temporal correlations as altered neural communication does not adequately consider how shifts in shared variance across networks may influence these measures without reflecting true biological reorganization.

The discussion does not substantively engage with the longstanding debate over whether sensory experience plays an instructive or permissive role in cortical development.

The relationship between resting-state and task-based findings in blindness remains unclear.

Reviewer #2 (Public review):

Summary:

Tian et al. explore the developmental origins of cortical reorganization in blindness. Previous work has found that a set of regions in the occipital cortex show different functional responses and patterns of functional correlations in blind vs. sighted adults. Here, Tian et al. explore how this organization arises over development. Is the "starting state" more like the blind pattern, or more like the adult pattern? Their analyses reveal that the answer depends on the particular networks investigated. Some functional connections in infants look more like blind than sighted adults; other functional connections look more like sighted than blind adults; and others fall somewhere in the middle, or show an altogether different pattern in infants compared with both sighted and blind adults.

Strengths:

The paper addresses very important questions about the starting state in the developing visual cortex, and how cortical networks are shaped by experience. Another clear strength lies in the unequivocal nature of many results. Many results have very large effect sizes, critical interactions between regions and groups are tested and found, and infant analyses are replicated in split halves of the data.

Weaknesses:

While potential roles of experience (e.g., visual, cross-modal) are discussed in detail, little consideration is given to the role of experience-independent maturation. The infants scanned are extremely young, only 2 weeks old. It is possible then that the sighted adult pattern may still emerge later in infancy or childhood, regardless of infant visual experience. If so, the blind adult pattern may depend on blindness-related experience only (which may or may not reflect "visual" experience per se). In short, it is not clear that birth, or the first couple weeks of life, are a clear cut "starting point" for development, after which all change can be attributed to experience.

Reviewer #3 (Public review):

Summary

This study aimed to investigate whether the differences observed in the organization of visual brain networks between blind and sighted adults result from a reorganization of an early functional architecture due to blindness, or whether the early architecture is immature at birth and requires visual experience to develop functional connections. This question was investigated through the comparison of 3 groups of subjects with resting-state functional MRI (rs-fMRI). Based on convincing analyses, the study suggests that: 1) secondary visual cortices showed higher connectivity to prefrontal cortical regions (PFC) than to non-visual sensory areas (S1/M1 and A1) in infants like in blind adults, in contrast to sighted adults; 2) the V1 connectivity pattern of infants lies between that of sighted adults (showing stronger functional connectivity with non-visual sensory areas than with PFC) and that of blind adults (showing stronger functional connectivity with PFC than with non-visual sensory areas); 3) the laterality of the connectivity patterns of infants resembled those of sighted adults more than those of blind adults, but infants showed a less differentiated fronto-occipital connectivity pattern than adults.

Strengths

- The question investigated in this article is important for understanding the mechanisms of plasticity during typical and impaired development, and the approach considered, which compares different groups of subjects including, neonates/infants and blind adults, is highly original.

- Overall, the presented analyses are solid and well detailed, and the results and discussion are convincing.

Weaknesses

- While it is informative to compare the "initial" state (close to birth) and the "final" states in blind and sighted adults to study the impact of post-natal and visual experience, this study does not analyze the chronology of this development and when the specialization of functional connections is completed. This would require investigating the evolution of functional connectivity of the visual system as a function of visual experience and thus as a function of age, at least during toddlerhood given the early and intense maturation of the visual system after birth. This could be achieved by analyzing different developmental periods using open databases such as the Baby Connectome Project.

- The rationale for grouping full-term neonates and preterm infants (scanned at term-equivalent age) is not understandable when seeking to perform comparisons with adults. Even if the study results do not show differences between full-terms and preterms in terms of functional connectivity differences between regions and of connectivity patterns, preterms group had different neurodevelopment and post-natal (including visual) experiences (even a few weeks might have an impact). And actually they show reduced connectivity strength systematically for all regions compared with full-terms (Sup Fig 7). Considering a more homogeneous group of neonates would have strengthen the study design.

- The rationale for presenting results on the connectivity of secondary visual cortices before the one of primary cortices (V1) could be clarified.

- The authors acknowledge the methodological difficulties for defining regions of interest (ROIs) in infants in a similar way as adults. Since the brain development is not homogeneous and synchronous across brain regions (in particular with the frontal and parietal lobes showing a delayed growth), this poses major problems for registration. This raises the question of whether the study findings could be biased by differences in ROI positioning across groups.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

The authors are appropriately cautious in many parts of the discussion and include several helpful control analyses. Nonetheless, additional clarification of key assumptions and potential confounds would strengthen the paper.

(1) The current framing labels vision as "instructive" and blindness as "reorganizing," but it is unclear why these two experiential factors are characterized differently. Both involve activity-dependent changes to functional architecture from a shared immature scaffold. Labeling them differently risks conflating divergent outcomes with distinct underlying mechanisms. Just because visual and blind adults show different patterns of functional connectivity does not mean they reflect separate processes. While the discussion briefly acknowledges the possibility of shared plasticity mechanisms, much of the framing across the manuscript, including in the abstract and introduction, implies a dichotomy. A clearer articulation of the criteria used to assign these labels, or reconsideration of whether such a distinction is warranted, would improve conceptual clarity. The current framing appears analogous to saying that "heat causes expansion" and "cold causes contraction" as if these were separate mechanisms, when they are actually two directions of change along a single factor: temperature. A more parsimonious framework, such as activity-dependent reweighting of pre-existing connectivity, may better capture the nature of plasticity at play in both sighted and blind development.

Following the reviewer’s suggestion, we have revised the manuscript to clarify that both vision and blindness can be understood as manifestations of a common framework of experience-driven plasticity. We removed all mention of reorganization and clarify and modified the wording throughout.

Specifically:

Abstract: “Are infant visual cortices functionally like those of sighted adults, with blindness leading to functional change? We find that, on the contrary that secondary visual cortices of infants are functionally more like those of blind adults: stronger coupling with PFC than with nonvisual sensory-motor networks, suggesting that visual experience modifies elements of the sighted-adult long-range functional connectivity profile. Infant primary visual cortices are in-between blind and sighted adults i.e., more balanced PFC and sensory-motor connectivity than either adult group. The lateralization of occipital-to-frontal connectivity in infants resembles the sighted adults, consistent with the idea that blindness leads to functional change. These results suggest that both vision and blindness modify functional connectivity through experience-driven (i.e., activity-dependent) plasticity.” (Page 1, Line 13)

Introduction: We replaced “blindness leads to functional reorganization” with “blindness modifies this functional connectivity” (Page 2, Line 52), and the following sentence has also been modified to: “lifetime visual experience shapes connectivity toward the sighted-adult pattern” (Page 2, Line 54) For the lateralization patterns, we now describe them as “blindness-related modification” rather than “reorganization”, to keep the interpretation descriptive rather than mechanistic. (Page 4, Line 114),

(2) In interpreting the functional correlation differences, the discussion should more explicitly consider how statistical interdependence between areas could influence the observed results. For example, an increase in shared variance between visual and motor areas, such as might result from visually guided action, could result in a reduction in the apparent strength of visual-prefrontal temporal correlation (at the resolution of fMRI) without any true biological change in communication between visual-prefrontal cortex. This possibility is not ruled out by reporting groupwise patterns of relative connectivity. A more cautious systems-level framing could help clarify the distinction between neural plasticity and statistical redistribution of variance.

We thank the reviewer for raising this important point. We agree that resting-state fMRI provides a measure of statistical synchrony in BOLD signals rather than direct causal interactions between regions. This a fundamental limitation of resting state fMRI, which we now note in the Discussion section. Such changes in correlation are consistent with a variety of underlying biological mechanisms. Online task is one factor that influences cross-region correlations. In the current study, both blind and sighted groups were measured while blindfolded and were not performing visually guided actions during the resting state fMRI scans. It is possible that past visual-guided action experience changes the resting state correlations of sighted participants. Indeed, this is one interesting hypothesis.

In the revised Discussion, we now explicitly note this limitation and clarify that differences in FC do not by themselves establish whether or how underlying neurophysiological mechanisms are changed. We also emphasize that future work will need to investigate whether FC changes are accompanied by alterations in structural connectivity and to probe causal interactions and mechanistic underpinnings as follows：

“Resting-state functional connectivity captures synchrony in BOLD signal fluctuations rather than causal interactions and differences in functional connectivity cannot on their own reveal how underlying neurophysiological mechanisms are modified.” (page 13,line 342)

“Future studies will be needed to determine whether these functional changes are accompanied by alterations in structural connectivity, and to probe causal interactions and mechanistic underpinnings.” (page 13,line 350)

(3) The mechanistic interpretation of group differences in visual-motor coupling would benefit from stronger network-level justification. Direct connections between these areas are sparse in primates. If effects reflect indirect polysynaptic interactions or shared thalamic input, as the authors suggest, one might expect corresponding group differences in intermediate regions (e.g., parietal cortex, thalamus) that mediate these interactions. Is there any evidence for this in the data?

We thank the reviewer for raising this point. We agree and as noted above, resting state fMRI cannot distinguish between direct causal interactions between two regions and ones that a mediating region is involved. This is a fundamental limitation of resting state fMRI. The current study further focused on testing a specific hypothesis motivated by previously observed group differences between blind and sighted adults and our analyses focused on ROI-to-ROI connectivity between occipital, frontal, and sensory-motor cortices, and did not include these additional regions. In prior work, we and others, have looked at effects in parietal cortices (Abboud & Cohen, 2019; Bedny et al., 2009; Deen et al., 2015; Kanjlia et al., 2016, 2021; Sen et al., 2022). In blindness, parietal networks show increased correlations with some visual areas, rather than decreased. Regarding the thalamus, there is less clear evidence and there is some ongoing work trying to address this question. A couple of studies suggest that there is indeed increased connectivity between some parts of the thalamus and visual cortex in blindness. Although the anatomical information is limited, some of the work suggests that this increase is with higher-cognitive nuclei of the thalamus (Bedny et al., 2011; Liu et al., 2007).

We agree that this is an important direction for future work. To acknowledge this point, we have revised the manuscript to highlight the potential role of cortical and subcortical hub regions in mediating connectivity changes. The text has been modified as follows:

“Connectivity changes between two areas could be mediated by ‘third-party’ hub regions. For example, posterior parietal cortex serves as a cortical hub for multisensory integration and visuo-motor coordination and could mediate occipital-to-sensory-motor communication (Rolls et al., 2023; Sereno & Huang, 2014). Subcortical structures such as the thalamus could also play a mediating role (Vega-Zuniga et al., 2025).” (page 13,line 345)

(4) The discussion would benefit from deeper engagement with prior work on experience-dependent plasticity, particularly the longstanding distinction between instructive and permissive roles of experience. While the authors briefly define these concepts and reference their historical use, a more explicit consideration of how their findings relate to this broader literature would help clarify whether such distinctions are necessary or appropriate.

We thank the reviewer for this thoughtful suggestion to engage more explicitly with the longstanding literature on instructive versus permissive roles of experience. However, most of this literature comes from animal models, where experimental manipulations of the anatomical structure, of experience itself (e.g., controlled rearing studies) and sometimes of neural activity patterns allow clear tests of these mechanisms. Such manipulations are not feasible in humans. The terminology in the animal literature does not directly map onto the methods and data available in the present study or in other work with humans. For this reason, the current data does not allow us to fully engage with the debates in the animal literature and doing risks overinterpreting our findings.

Nevertheless, we agree that once the instructive/permissive framework has been introduced, it is important to clarify how our results relate to it, rather than only providing definitions. We have therefore added the following text to the discussion:

“In humans, such manipulations are not feasible, leaving us to study only the consequences of the presence or absence of vision. Under an instructive account, visual and multisensory experience could strengthen coupling between visual and other non-visual sensory-motor cortices through coordinated activity, thereby establishing the sighted-adult connectivity pattern. In the absence of visual input, by contrast, the lack of such coordinated activity may prevent these couplings from being established. Alternatively, vision may act permissively, indirectly enabling maturational processes that shift connectivity toward the sighted-adult configuration.” (page 14,line 362)

(5) The revised discussion acknowledges the divergence between resting-state and task-based findings, but does not fully frame the theoretical implications of this discrepancy. Although this study cannot resolve the issue with its own data, a more integrative discussion could help clarify whether these measures reflect distinct functional states, developmental trajectories, or mechanisms of plasticity. Without such framing, readers are left without clear guidance on how to reconcile the present results with prior work on cross-modal recruitment in blindness.

We thank the reviewer for this thoughtful comment. We agree that know how resting-state evidence relates to task-based evidence is a fundamentally important issue. We now discuss this more in the Introduction as well as in the Discussion.

There is a sizable literature of both task-based and resting state studies. Some of prior studies have measured resting state and task-based data within the same participants and found relationships (Kanjlia et al., 2016, 2021; Lane et al., 2015). We now clarify this in the introduction. These studies find that within visual cortices of blind people, the task-based profile of a cortical area is related to its resting state connectivity pattern (Abboud & Cohen, 2019; Deen et al., 2015; Kanjlia et al., 2016, 2021). This suggests that these two measures are related. However, the timecourse of this relationship, the developmental trajectory and mechanism of plasticity is not known. We note this now in the introduction on page 2. Primarily this is because there is very little relevant developmental evidence. For example, in the current study we find that the resting state profile of secondary visual networks in infants is similar to that of blind adults. However, we do not know whether the visual cortices of infants show task-based cross modal responses. To our knowledge nobody has tested this question. We agree with the reviewer that raising this question in the paper is better than not commenting on the relationship at all.

To address the reviewer’s comment, we have expanded the discussion to situate our results within a developmental framework, highlighting how early intrinsic connectivity may scaffold alternative trajectories shaped by either visual experience or blindness. The revised text now reads as follows:

“Conversely, for people who remain blind throughout life, visual-PFC connectivity could enable recruitment of visual cortices for higher-order non-visual functions, such as language and executive control (Bedny et al., 2011; Kanjlia et al., 2021). Our results suggest that blind adults may build on connectivity patterns already present in infancy: like blind adults, sighted infants show stronger occipital–PFC than occipital–sensory–motor coupling. Repeated engagement of occipital networks during higher cognitive tasks in early development could intern enhance connectivity and specialization of visual networks for non-visual higher-order functions.

Some prior studies have measured resting-state and task-based functional profiles in the same participants. These studies find that within visual cortices of blind people, the task-based profile of a cortical area is related to its resting state connectivity pattern (citations.) This suggests that these two measures are related. However, the timecourse of this relationship, the developmental trajectory and mechanism of plasticity is not known. Primarily this is because there is very little relevant developmental evidence. For example, in the current study we find that the resting state profile of secondary visual networks in infants is similar to that of blind adults. However, we do not know whether the visual cortices of infants show enhanced task-based cross modal responses, relative to sighted adults and how this compares to responses observed in blind adults. Future work with infants and children would be able to address this question.

In the current study, the clearest evidence for functional change driven by blindness was observed for laterality. Connectivity lateralization in sighted infants resembles that of sighted adults, in both V1 and secondary visual cortices. Relative to both sighted infants and sighted adults, blind adults show more lateralized connectivity patterns between occipital and prefrontal cortices. Previous studies suggest that in people born blind occipital and non-occipital language responses are co-lateralized (Lane et al., 2017; Tian et al., 2023). We speculate that habitual activation of visual cortices by higher-cognitive tasks, such as language, which are themselves highly lateralized, contributes to this biased connectivity pattern of occipital cortex in blindness. Taken together, these results suggest a developmental framework in which intrinsic connectivity present in infancy provides a scaffold that is subsequently shaped and reinforced by experience-dependent recruitment, through either visual experience or the lifelong absence of vision in blindness. Longitudinal work across successive developmental stages will be crucial to test how the alternative trajectories shaped by visual experience versus blindness unfold over development.” (page 14-15)

(6) The split-half reliability analysis is a valuable control. Additional details would clarify what these noise ceilings reflect. Were the rsFC patterns for each ROI calculated only for the ROIs included in the current study or was a broader assessment across the whole brain performed? It also would be helpful to report whether reliability differed for individual ROIs within and between groups. Even if global reliability is matched, selective differences could influence group comparisons. Several infants in the dhcp dataset were scanned twice. Were any second scans included in the current analyses? Comparing first versus second scans directly could strengthen the claim that several weeks of visual experience are insufficient to shift connectivity toward a sighted adult profile.

Thanks to the reviewer’s comments on the reliability of the current study.

In the present study, the noise ceiling was computed from the reliability of the ROI-wise FC profiles used across all analyses. Reliability was estimated using a split-half procedure: each rs-fMRI time series was divided into two equal halves, FC among all ROIs included in the study was computed separately for each half, and the noise ceiling for each ROI was defined as the Pearson correlation between its two FC profiles. Then we averaged these ROI-wise noise ceilings to evaluate group-level reliability, which exceeded 0.70 in all three groups and found no significant difference across groups. This provides an estimate of the upper bound on explainable variance for the exact FC features subjected to statistical testing (Lage-Castellanos et al., 2019). A brief description has been added to the manuscript (page 19, line 518).

Regarding the reviewer’s question about the scope of rsFC features used in the noise-ceiling analysis: we computed noise ceilings only for the ROIs included in the present study, because all analyses in this work were conducted at the ROI–ROI level and did not involve voxelwise whole-brain FC. Thus, the noise-ceiling estimates correspond directly to the full set of FC features on which all statistical comparisons were based.

As suggested by the reviewer, we examined noise ceilings for each ROI separately. All ROIs showed high absolute reliability (noise ceiling > 0.80) across the three groups, indicating that the ROI-wise FC estimates are generally robust across participants. Although many ROIs exhibited statistically significant group differences in noise ceiling (one-way ANOVA, p < 0.05), the effect sizes were small to moderate (partial η2 < 0.14). These differences indicate that reliability may vary modestly across groups at the ROI level, and we cannot fully determine whether such variability contributes to the observed different FC patterns across groups. We have included this point in the revised manuscript (page 19, line 525), along with the full statistical results for the ROI-wise noise ceilings in the Supplementary Table S2.

Last, we fully agree that longitudinal comparisons across multiple time points can provide important insights into how early visual experience shapes connectivity. At the same time, in the present dataset, the first scan occurred at a preterm age and the second at term-equivalent age. The differences between the first and second scans would reflect not only additional weeks of visual input, but also differences in prematurity status and overall neurodevelopmental maturity, which would make the interpretation of such comparisons difficult in the context of our current aims. We have clarified in the revised manuscript that only term-equivalent (second) scans were included. We see careful longitudinal work as an important avenue for addressing this question more directly.

(7) The signal dropout assessment in the infant dataset is a valuable quality control step. Applying the same metric to the adult datasets would help harmonize preprocessing across groups and increase confidence in group-level comparisons.

Thank you for this valuable suggestion. Following your comment, we applied the same signal dropout assessment to the adult datasets. One participant in the sighted adult group and two participants in the blind adult group showed signal dropout in one ROI each. The corresponding results are now included in the Supplementary Materials (Figure S13). The findings remain unchanged after this additional control analysis. We also add the relevant content in the Method part as follows:

“The same signal dropout assessment was also applied to the blind and sighted adults to ensure consistent quality control across groups. One participant in the sighted adult group and two participants in the blind adult group exhibited signal dropout in one ROI each. Excluding these participants did not alter the group-level results (see Figure S13).” (page 16, line 449)

Minor:

(8) The authors added accurate anatomical descriptions to the methods but a less precise characterization remains in the introduction: "Anatomically, these regions correspond roughly to the location of areas such as motion area V5/MT+, the lateral occipital complex (LO), V3a and V4v in sighted people."

We thank the reviewer for this helpful comment. We have revised the Introduction to provide a fuller anatomical description, consistent with the Methods. The text now reads:

“Anatomically, these regions in sighted people approximately correspond to the locations of motion-sensitive V5/MT+ and the lateral occipital complex (LO), as well as ventral portions of occipito-temporal cortex including V4v and dorsal portions including V3a. The occipital ROI also extends ventrally into the middle portion of the ventral temporal lobe and dorsally into the intraparietal sulcus and superior parietal lobule.” (page 3, line 88)

(9)Typo: "lager effect" should be "larger effect."

Secondary visual cortices showed a significant within > between difference in both groups, with a lager effect in the blind group (post-hoc tests, Bonferroni-corrected paired: t-test: sighted adults within hemisphere > between hemisphere: t (49) = 7.441, p = 0.012; blind adults within hemisphere > between hemisphere: t (29) = 10.735, p < 0.001; V1: F(1, 78) =87.211, p < 0.001).

We thank the reviewer for catching this typo. We have corrected “lager effect” to “larger effect” in the revised manuscript. (page 9, line 214)

Reviewer #2 (Recommendations for the authors):

All of my other concerns were adequately addressed.

We thank the reviewer for their positive evaluation, and we are glad that our revisions have addressed their concerns.

Reviewer #3 (Recommendations for the authors):

In my view, qualifying infants as "sighted" is confusing and unnecessary: why not simplifying and homogenizing the wording along the manuscript and figures?

We thank the reviewer for this suggestion. We agree and have revised the manuscript to use consistent wording, avoiding the qualification of infants as “sighted.”

l188, I don't understand the sentence "By contrast, in sighted adults, this cross-hemisphere difference is weak or absent."

We thank the reviewer for noting that this sentence was unclear. We have revised the text to provide a more precise explanation. The text now reads:

“By contrast, in sighted adults this lateralized pattern is weaker: visual areas in each hemisphere show only a modest preference for ipsilateral prefrontal cortices, and connectivity with the contralateral PFC remains comparatively strong.” (page 8, line 207)

l193: "Secondary visual cortices showed a significant within > between difference in both groups, with a lager effect in the blind group": providing effect sizes for the 2 groups would strengthen this result (+ note the typo laRger).

- Figure S7, S11: Please add titles of y-axes.

Thank you for this helpful suggestion. We have corrected the typo and added the effect sizes for both groups in the revised text. The revised sentence now reads as follows:

“Secondary visual cortices showed a significant within > between difference in both groups, with a larger effect in the blind group (post-hoc tests, Bonferroni-corrected paired: t-test: sighted adults within hemisphere > between hemisphere: t (49) = 7.441, p = 0.012, cohen’d = 0.817; blind adults within hemisphere > between hemisphere: t (29) = 10.735, p < 0.001, cohen’d = 1.96).” (page 9, line 214)

Titles of the y-axes have also been added to Figures S7 and S11.

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Abstract

Earlier version: Comparisons of visual cortex function across blind and sighted adults reveals effects of experience on human brain function. Since almost all research has been done with adults, little is known about the developmental origins of plasticity. We compared resting state functional connectivity of visual cortices of blind adults (n = 30), blindfolded sighted adults (n = 50) to a large cohort of sighted infants (Developing Human Connectome Project, n = 475). Visual cortices of sighted adults show stronger coupling with non-visual sensory-motor networks (auditory, somatosensory/motor), than with higher-cognitive prefrontal cortices (PFC). In contrast, visual cortices of blind adults show stronger coupling with higher-cognitive PFC than with nonvisual sensory-motor networks. Are infant visual cortices functionally like those of sighted adults? Alternatively, do infants start like blind adults, with vision required to set up the sighted adult pattern? Remarkably, we find that, in infants, secondary visual cortices are more like those of blind adults: stronger coupling with PFC than with nonvisual sensory-motor networks, suggesting that visual experience establishes elements of the sighted-adult long-range connectivity. Infant primary visual cortices are in-between blind and sighted adults i.e., equal PFC and sensory-motor connectivity. The lateralization of occipital-to-frontal connectivity in infants resembles the sighted adults, consistent with reorganization by blindness. These results reveal instructive effects of vision and reorganizing effects of blindness on functional connectivity.

Revised version: Comparisons of visual cortex function across blind and sighted adults reveals effects of experience on human brain function. Since almost all research has been done with adults, little is known about the developmental origins of plasticity. We compared resting state functional connectivity of visual cortices of blind adults (n = 30), blindfolded sighted adults (n = 50) to a large cohort of infants (Developing Human Connectome Project, n = 475). Visual cortices of sighted adults show stronger coupling with non-visual sensory-motor networks (auditory, somatosensory/motor), than with higher-cognitive prefrontal cortices (PFC). In contrast, visual cortices of blind adults show stronger coupling with higher-cognitive PFC than with nonvisual sensory-motor networks. Are infant visual cortices functionally like those of sighted adults, with blindness leading to functional change? We find that, on the contrary that secondary visual cortices of infants are functionally more like those of blind adults: stronger coupling with PFC than with nonvisual sensory-motor networks, suggesting that visual experience modifies elements of the sighted-adult long-range functional connectivity profile. Infant primary visual cortices are in-between blind and sighted adults i.e., more balanced PFC and sensory-motor connectivity than either adult group. The lateralization of occipital-to-frontal connectivity in infants resembles the sighted adults, consistent with the idea that blindness leads to functional change. These results suggest that both vision and blindness modify functional connectivity through experience-driven (i.e., activity-dependent) plasticity.

Introduction

Relative to sighted adults, visual cortices of adults born blind show enhanced responses during non-visual tasks, such as reading braille and localizing sounds as well as distinctive patterns of long-range functional connectivity with non-visual networks (Abboud & Cohen, 2019; Bedny et al., 2011; Burton et al., 2012, 2014; Butt et al., 2013; Collignon et al., 2011; Deen et al., 2015; Kanjlia et al., 2016, 2021; Lane et al., 2015; Liu et al., 2007; Sadato et al., 1996; Striem-Amit et al., 2015; Watkins et al., 2012). Since almost all research thus far has been done with adults, a key outstanding question concerns the developmental origins of these experience-based differences.

Earlier version: One possibility is that at birth, infant visual cortices start out in the ‘prepared’ sighted adult state and blindness leads to functional reorganization. Alternatively, infant visual cortices may start out functionally similar to those of blind adults and lifetime visual experience establishes the sighted adult pattern. To distinguish between these possibilities, we compare the function of visual cortices across blind adults, sighted adults and a large cohort of 2-week-old sighted infants (Developing Human Connectome Project, dHCP, n = 475). We examined functional connectivity using resting state data, which provide a common measure of cortical function across these diverse populations.

Revised version: One possibility is that at birth, infant visual cortices start out in the ‘prepared’ sighted adult state and blindness modifies this functional connectivity. Alternatively, infant visual cortices may start out functionally similar to those of blind adults and lifetime visual experience shapes connectivity toward the sighted-adult pattern. To distinguish between these possibilities, we compare the long-range function connectivity of visual cortices across blind adults, sighted adults and a large cohort of 2-week-old infants (Developing Human Connectome Project, dHCP, n = 475). Using resting state data provide a common measure of cortical function across these diverse populations.

Earlier version: To our knowledge no prior studies have compared infants to multiple populations of adults with different sensory experiences. Previous studies comparing sighted infants to sighted adults have largely reported similarity across groups (Barttfeld et al., 2018; Doria et al., 2010; Fransson et al., 2009; Gao et al., 2009; Liu et al., 2008; Zhang et al., 2019). However, these studies focused on whether large scale functional networks are present in infancy e.g., stronger connectivity of regions within the visual network, than between visual and auditory regions. Studies comparing blind and sighted adults find differences across groups in which non-visual networks are most strongly coupled with the visual system i.e., visual cortices of sighted adults show stronger coupling with non-visual sensory-motor networks (i.e., auditory, somatosensory/motor) than higher-cognitive systems; by contrast, in blind adults, visual cortex coupling is stronger with higher-cognitive prefrontal cortices (PFC) than with non-visual sensory-motor networks (Abboud & Cohen, 2019; Bedny et al., 2011; Burton et al., 2014; Deen et al., 2015; Kanjlia et al., 2021; Liu et al., 2007; Qin et al., 2013; Striem-Amit et al., 2015; Watkins et al., 2012; Yu et al., 2008). In the current study, we compare this experience-sensitive functional signature across infants, sighted and blind adults.

Revised version: To our knowledge no prior studies have compared infants to multiple populations of adults with different sensory experiences. Previous studies comparing infants to sighted adults have largely reported similarity across groups (Barttfeld et al., 2018; Doria et al., 2010; Fransson et al., 2009; Gao et al., 2009; Liu et al., 2008; Zhang et al., 2019). However, these studies focused on whether large scale functional networks are present in infancy e.g., stronger connectivity of regions within the visual network, than between visual and auditory regions. Studies comparing blind and sighted adults find differences across groups in which non-visual networks are most strongly coupled with the visual system i.e., visual cortices of sighted adults show stronger coupling with non-visual sensory-motor networks (i.e., auditory, somatosensory/motor) than higher-cognitive systems; by contrast, in blind adults, visual cortex coupling is stronger with higher-cognitive prefrontal cortices (PFC) than with non-visual sensory-motor networks (Abboud & Cohen, 2019; Bedny et al., 2011; Burton et al., 2014; Deen et al., 2015; Kanjlia et al., 2021; Liu et al., 2007; Qin et al., 2013; Striem-Amit et al., 2015; Watkins et al., 2012; Yu et al., 2008). In the current study, we compare this experience-sensitive functional signature across infants, sighted and blind adults.

Earlier version: We measured the connectivity profile of four occipital ‘visual’ areas that show cross-modal plasticity in blindness, i.e., respond to non-visual tasks in blind people: Three secondary visual areas (located in lateral, dorsal and parts of the ventral occipital cortex) and the primary visual cortex (V1).

Revised version: We measured the connectivity profile of four occipital ‘visual’ areas that show cross-modal plasticity in blindness, i.e., are active during non-visual tasks in blind people and show related changes in resting-state functional connectivity.

Earlier version: The three secondary visual areas have been found to respond to different non-visual tasks in blind people: language tasks, numerical reasoning tasks and executive control tasks respectively. Enhanced coupling with PFC in blindness is observed across all of three occipital regions. However, in blind people each region shows preferential coupling with distinct subregions of PFC with analogous functional profiles (e.g., language responsive occipital areas are more coupled with language responsive PFC) (Kanjlia et al., 2016, 2021; Lane et al., 2015). The precise visual functions of these secondary visual regions in sighted people are not known. Anatomically, these regions correspond roughly to the location of areas such as motion area V5/MT+, the lateral occipital complex (LO), V3a and V4v in sighted people (Tootell et al., 1997; Van Essen et al., 2001).

Revised version: We focused on three functionally-distinct secondary visual areas (located in lateral, dorsal and parts of the ventral occipital cortex) and the primary visual cortex (V1). The three secondary visual areas have been found to respond to different non-visual tasks in blind people: language tasks, numerical reasoning tasks and executive control tasks respectively. Enhanced coupling with PFC is observed across all of three occipital regions in blind adults. However, each region shows preferential coupling with a distinct subregion of PFC with analogous functional profiles i.e., language responsive occipital areas are more coupled with language responsive PFC (Kanjlia et al., 2016, 2021; Lane et al., 2015). These observations suggest that resting-state and task-based functional profiles are related, although the functional and developmental nature of this relationship remains an open question.

Added in the revised version: The precise visual functions of the studied secondary visual regions in sighted people are not known. Anatomically, these regions in sighted people approximately correspond to the locations of motion-sensitive V5/MT+ and the lateral occipital complex (LO), as well as ventral portions of occipito-temporal cortex including V4v and dorsal portions including V3a. The occipital ROI also extends ventrally into the middle portion of the ventral temporal lobe and dorsally into the intraparietal sulcus and superior parietal lobule (Tootell et al., 1997; Van Essen et al., 2001).

Earlier version: We also examined connectivity of anatomically defined primary visual cortex (V1), which likewise shows altered task-based responses and functional connectivity in congenitally blind adults (Amedi et al., 2003; Bedny et al., 2011; Burton et al., 2014; Butt et al., 2013; Lane et al., 2015; Raz et al., 2005; Sadato et al., 1996; Striem-Amit et al., 2015; Yu et al., 2008). Since many previous studies have found that blindness alters the balance of connectivity between visual cortex and higher-order prefrontal as opposed to sensory-motor regions, this was our primary outcome measure. We also examined changes in connectivity lateralization – i.e., the balance of between vs. within hemisphere connectivity.

Revised version: We also examined connectivity of anatomically defined primary visual cortex (V1), which likewise shows altered task-based responses and functional connectivity in congenitally blind adults (Amedi et al., 2003; Bedny et al., 2011; Burton et al., 2014; Butt et al., 2013; Lane et al., 2015; Raz et al., 2005; Sadato et al., 1996; Striem-Amit et al., 2015; Yu et al., 2008).

Added in the revised version: Since many previous studies have found that blindness alters the balance of connectivity between visual cortex and higher-order prefrontal as opposed to sensory-motor regions, this was our primary outcome measure (Abboud & Cohen, 2019a; Bedny et al., 2011; Burton et al., 2014; Deen et al., 2015; Heine et al., 2015; Kanjlia et al., 2016, 2021; Lane et al., 2015; Y. Liu et al., 2007; Sen et al., 2022; Striem-Amit et al., 2015). We also examined changes in connectivity lateralization – i.e., the balance of between vs. within hemisphere connectivity, since prior task-based studies have found laterality changes in blindness as well as co-lateralization of occipital and non-occipital networks in this population (Kanjlia et al., 2021; Lane et al., 2017; Tian et al., 2023).

Earlier version: To preview the results, we find that, in infants, the long-range functional connectivity profile of secondary visual areas resembles that of blind adults, whereas V1 falls between blind and sighted adult populations. Relative to sighted adults, both blind adults and infants show stronger coupling between visual cortices and PFC and weaker coupling between visual cortices and non-visual sensory-motor networks. This suggests that vision plays an instructive role in setting up the balance of connectivity between occipital cortex and non-visual networks. In contrast, connectivity lateralization patterns appear to reflect blindness-related reorganization.

Revised version: To preview the results, we find that, in infants, the long-range functional connectivity profile of secondary visual areas resembles that of blind adults, whereas V1 falls between blind and sighted adult populations. Relative to sighted adults, both blind adults and infants show stronger coupling between visual cortices and PFC and weaker coupling between visual cortices and non-visual sensory-motor networks. This suggests that vision contributes to modifying the balance of connectivity between occipital and non-visual networks after birth. In contrast, connectivity lateralization patterns appear to reflect blindness-related modification.

Results

Earlier version: Connectivity profile of secondary visual cortices in sighted infants is more similar to that of blind than sighted adults

Revised version: Connectivity profile of secondary visual cortices in infants is more similar to that of blind than sighted adults

We first examined the long-range functional connectivity of (three) secondary visual areas with sensory-motor areas on the one hand, and higher-order PFC networks on the other. In sighted adults, all three secondary visual areas showed stronger functional connectivity with non-visual sensory-motor areas (primary somatosensory and motor cortex, S1/M1, and primary auditory cortex, A1) than with higher-cognitive prefrontal cortices (PFC). By contrast, in blind adults, all secondary visual regions showed stronger functional connectivity with PFC than with non-visual sensory-motor areas (S1/M1 and A1) (group (sighted adults, blind adults) by ROI (PFC, non-visual sensory) interaction effect: F(1, 78) = 148.819, p < 0.001; post-hoc Bonferroni-corrected paired t-test, sighted adults: non-visual sensory > PFC: t (49) = 9.722, p < 0.001; blind adults: non-visual sensory < PFC: t (29) =8.852, p < 0.001; Fig. 1).

Functional connectivity of secondary visual cortices.

Earlier version: (A) Bar graph shows functional connectivity (r) of secondary visual cortices (blue) to non-visual sensory motor areas (purple) and prefrontal cortices (green), averaged across occipital, PFC and sensory-motor ROIs (A1 and S1/M1) in sighted adults, blind adults and sighted infants. Regions of interest (ROI) displayed on the left. Note that regions extend to ventral surface, not shown. See Supplementary Figure S5 for the full views of three occipital ROIs. (B) Circle plots represent the connectivity of secondary visual cortices to non-visual networks, min-max normalized to [0,1], i.e., as a proportion. OC: occipital cortices; MTH: math-responsive region; LG: language-responsive region; EF: executive function-responsive (response-conflict) region. Asterisks (*) denote significant Bonferroni-corrected pairwise comparisons (p < .05).

Revised version: (A) Violin plots show the distributions of functional connectivity (r) of secondary visual cortices (blue) to non-visual sensory-motor areas (purple) and prefrontal cortices (green), averaged across three occipital, PFC, and sensory-motor ROIs (A1 and S1/M1) in sighted adults, blind adults, and infants. Individual dots denote mean connectivity values per participant. Grey trend lines illustrate within-participant changes across sensory–motor and prefrontal targets. Dark-blue horizontal markers indicate group averages. Regions of interest (ROI) displayed on the left. Note that regions extend to ventral surface, not shown. See Supplementary Figure S5 for the full views of three occipital ROIs. (B) Circle plots represent the connectivity of secondary visual cortices to non-visual networks, min-max normalized to [0,1], i.e., as a proportion. OC: occipital cortices; MTH: math-responsive region; LG: language-responsive region; EF: executive function-responsive (response-conflict) region. Asterisks (*) denote significant Bonferroni-corrected pairwise comparisons (p < .05).

Earlier version: Like in blind adults, in sighted infants, secondary visual areas showed higher connectivity to PFC than to non-visual sensory-motor areas (S1/M1 and A1) (non-visual sensory < PFC paired t-test, t (474) = 20.144, p < 0.001) (Fig. 1). The connectivity matrix of sighted infants was more correlated with that of blind than sighted adults, but strongly correlated with both adult groups (secondary visual, PFC and non-visual sensory areas: sighted infants correlated to blind adults: r = 0.721, p < 0.001; to sighted adults: r = 0.524, p < 0.001; difference between correlations of infants to blind vs. to sighted adults: z = 3.77, p < 0.001; see Supplementary Figure S1 for the connectivity matrices).

Revised version: Like in blind adults, in infants, secondary visual areas showed higher connectivity to PFC than to non-visual sensory-motor areas (S1/M1 and A1) (non-visual sensory < PFC paired t-test, t (474) = 20.144, p < 0.001) (Fig. 1). The connectivity matrix of infants was more correlated with that of blind than sighted adults, but strongly correlated with both adult groups (secondary visual, PFC and non-visual sensory areas: infants correlated to blind adults: r = 0.721, p < 0.001; to sighted adults: r = 0.524, p < 0.001; difference between correlations of infants to blind vs. to sighted adults: z = 3.77, p < 0.001; see Supplementary Figure S1 for the connectivity matrices).

These results suggests that vision is required to set up elements of the sighted adult functional connectivity pattern, i.e., vision enhances occipital cortex connectivity to non-visual sensory-motor networks and dampens connectivity to higher-cognitive prefrontal networks.

We checked the robustness of these results in a number of ways. We first compared the effects across the three secondary visual regions and observed the same pattern across all. (Supplementary results and Supplementary Figure S2). Next, to check the robustness of the findings in infants we randomly split the infant dataset into two halves and did split-half cross-validation. Across all comparisons the results of the two halves were highly similar, suggesting the effects are robust (see Supplementary Figure S3 and S4). We performed this validation procedure for all analyses reported below with similar results.

The connectivity pattern of V1 is influenced by early visual experience and blindness

V1 showed the same dissociation between sighted and blind adults as secondary visual areas: in sighted adults, V1 has stronger functional connectivity with non-visual sensory-motor areas than with PFC. By contrast, in blind adults, V1 shows stronger connectivity with PFC than with non-visual sensory areas (group (sighted adults, blind adults) by ROI (PFC, non-visual sensory) interaction: F(1, 78) = 125.775, p < 0.001; post-hoc Bonferroni-corrected paired t-test, sighted adults non-visual sensory > PFC: t (49) = 9.404, p < 0.001; blind adults non-visual sensory < PFC: t (29) =7.128, p < 0.001; Fig. 2).

Functional connectivity of primary visual cortices (V1).

Earlier version: Regions of interest (ROI) displayed on the upper. Bar graph shows functional connectivity (r) of V1 to non-visual sensory motor areas (purple) and prefrontal cortices (green), averaged across three PFC ROIs and sensory-motor ROIs (S1/M1 and A1). Asterisks (*) denote significant Bonferroni-corrected pairwise comparisons (p < .05). Cross (†) denote marginal difference in Bonferroni-corrected pairwise comparisons (0.05< p <0.1).

Revised version: Violin plots show the distributions of functional connectivity (r) of V1 to non-visual sensory-motor areas (purple) and prefrontal cortices (green), averaged across three PFC ROIs and sensory-motor ROIs (S1/M1 and A1). In sighted adults, blind adults, and infants. Individual dots denote mean connectivity values per participant. Grey trend lines illustrate within-participant changes across sensory–motor and prefrontal targets. Dark-blue horizontal markers indicate group averages. Asterisks (*) denote significant Bonferroni-corrected pairwise comparisons (p < .05). Cross (†) denote marginal difference in Bonferroni-corrected pairwise comparisons (0.05< p <0.1).

Earlier version: The pattern for sighted infants in V1 fell between that of sighted and blind adults. The connectivity matrix of sighted infants (V1, PFC, and non-visual sensory) was equally correlated with blind and sighted adults (infants correlated to blind adults: r = 0.654, p < 0.001; to sighted adults: r = 0.594, p < 0.001; correlation of infants with blind vs. with sighted adults: z = 0.832, p = 0.406; see Supplementary Figure S1 for the connectivity matrices). The difference in connectivity strength between V1 to PFC and V1 to non-visual sensory regions was weaker in sighted infants than in sighted or blind adults (group (sighted adults, infants) by ROI (PFC, non-visual sensory) interaction effect: F(1, 523) = 92.21, p < 0.001; group (blind adults, infants) by ROI (PFC, non-visual sensory) interaction effect: F(1, 503) = 57.444, p < 0.001). V1 of sighted infants showed marginally stronger connectivity to non-visual sensory regions (S1/M1 and A1) than PFC (non-visual sensory regions > PFC, paired t-test, t (474) = 1.95, p = 0.052; Fig.2).

Revised version: The pattern for infants in V1 fell between that of sighted and blind adults. The connectivity matrix of infants (V1, PFC, and non-visual sensory) was equally correlated with blind and sighted adults (infants correlated to blind adults: r = 0.654, p < 0.001; to sighted adults: r = 0.594, p < 0.001; correlation of infants with blind vs. with sighted adults: z = 0.832, p = 0.406; see Supplementary Figure S1 for the connectivity matrices). The difference in connectivity strength between V1 to PFC and V1 to non-visual sensory regions was weaker in infants than in sighted or blind adults (group (sighted adults, infants) by ROI (PFC, non-visual sensory) interaction effect: F(1, 523) = 92.21, p < 0.001; group (blind adults, infants) by ROI (PFC, non-visual sensory) interaction effect: F(1, 503) = 57.444, p < 0.001). V1 of infants showed marginally stronger connectivity to non-visual sensory regions (S1/M1 and A1) than PFC (non-visual sensory regions > PFC, paired t-test, t (474) = 1.95, p = 0.052; Fig.2).

The dHCP cohort included both full-term neonates and preterm infants, scanned at their equivalent gestational age. Visual exposure therefore varied somewhat in duration across infants (from 0 to 19.71 weeks), with slightly longer exposure in preterm babies. This variation did not affect connectivity patterns either in V1 or secondary visual cortices (V1: r = 0.06, p = 0.192; secondary visual: r = 0.004, p = 0.923; see Supplementary Figure S6). We also compared the connectivity patterns of preterm (n = 90) and full-term infants (n = 385) and found no difference from each other or from the all-infant dataset (see Supplementary Figure S7). A few weeks of vision after birth is therefore insufficient to influence connectivity.

Earlier version: Evidence for blindness-related reorganization in laterality of occipito-frontal connectivity

Revised version: Evidence for blindness-related functional change in laterality of occipito-frontal connectivity

Earlier version: Compared to sighted adults, blind adults exhibit a stronger dominance of within-hemisphere connectivity over between-hemisphere connectivity. That is, in people born blind, left visual networks are more strongly connected to left PFC, whereas right visual networks are more strongly connected to right PFC. By contrast, in sighted adults, this cross-hemisphere difference is weak or absent. This difference between adult groups is observed for both V1 and secondary visual cortices (group (blind adults, sighted adults) by lateralization (within hemisphere, between hemisphere) interaction in secondary visual cortices: F(1, 78) = 131.51, p < 0.001). Secondary visual cortices showed a significant within > between difference in both groups, with a lager effect in the blind group (post-hoc tests, Bonferroni-corrected paired: t-test: sighted adults within hemisphere > between hemisphere: t (49) = 7.441, p = 0.012; blind adults within hemisphere > between hemisphere: t (29) = 10.735, p < 0.001; V1: F(1, 78) = 87.211, p < 0.001). In V1, only the blind group showed a significant within > between hemisphere effect (post-hoc Bonferroni-corrected paired: t-test: sighted adults within hemisphere < between hemisphere: t (49) = 3.251, p = 0.101; blind adults within hemisphere > between hemisphere: t (29) = 7.019, p < 0.001).

Revised version: Compared to sighted adults, blind adults exhibit a stronger dominance of within-hemisphere connectivity over between-hemisphere connectivity. That is, in people born blind, left visual networks are more strongly connected to left PFC, whereas right visual networks are more strongly connected to right PFC. By contrast, in sighted adults this lateralized pattern is weaker: visual areas in each hemisphere show only a modest preference for ipsilateral prefrontal cortices, and connectivity with the contralateral PFC remains comparatively strong. This difference between adult groups is observed for both V1 and secondary visual cortices (group (blind adults, sighted adults) by lateralization (within hemisphere, between hemisphere) interaction in secondary visual cortices: F(1, 78) = 131.51, p < 0.001;V1: F(1, 78) = 87.211, p < 0.001). Secondary visual cortices showed a significant within > between difference in both groups, with a larger effect in the blind group (post-hoc tests, Bonferroni-corrected paired: t-test: sighted adults within hemisphere > between hemisphere: t (49) = 7.441, p = 0.012, cohen’d = 0.817; blind adults within hemisphere > between hemisphere: t (29) = 10.735, p < 0.001, cohen’d = 1.96). In V1, only the blind group showed a significant within > between hemisphere effect (post-hoc Bonferroni-corrected paired: t-test: sighted adults within hemisphere < between hemisphere: t (49) = 3.251, p = 0.101; blind adults within hemisphere > between hemisphere: t (29) = 7.019, p < 0.001).

Earlier version: With respect to laterality, infants resembled sighted more than blind adults (Fig. 3). For secondary visual cortices, there was a significant difference between blind adults and sighted infants and no difference between sighted adults and sighted infants (group (blind adults, infants) by lateralization (within hemisphere, between hemisphere) interaction effect: F(1, 503) = 303.04, p < 0.001); (group (sighted adults, infants) by lateralization (within hemisphere, between hemisphere) interaction effect: F(1, 523) = 2.244, p = 0.135). A similar group by laterality interaction were observed for V1 (group (blind adults, infants) by lateralization (within hemisphere, between hemisphere) interaction: F(1, 503) = 123.608, p < 0.001; group (sighted adults, infants) by lateralization (within hemisphere, between hemisphere) interaction effect: F(1, 523) = 2.827, p = 0.093). This suggests that the enhancement of within over between hemisphere long-range connectivity is related to blindness-driven reorganization.

Revised version: With respect to laterality, infants resembled sighted more than blind adults (Fig. 3). For secondary visual cortices, there was a significant difference between blind adults and infants and no difference between sighted adults and infants (group (blind adults, infants) by lateralization (within hemisphere, between hemisphere) interaction effect: F(1, 503) = 303.04, p < 0.001); (group (sighted adults, infants) by lateralization (within hemisphere, between hemisphere) interaction effect: F(1, 523) = 2.244, p = 0.135). A similar group by laterality interaction were observed for V1 (group (blind adults, infants) by lateralization (within hemisphere, between hemisphere) interaction: F(1, 503) = 123.608, p < 0.001; group (sighted adults, infants) by lateralization (within hemisphere, between hemisphere) interaction effect: F(1, 523) = 2.827, p = 0.093). This suggests that the enhancement of within over between hemisphere long-range connectivity is related to blindness-driven functional change.

Within hemisphere vs. between hemisphere functional connectivity.

Earlier version: Bar graph shows within hemisphere (blue) and between hemisphere (orange) functional connectivity (r coefficient of resting state correlations) of secondary visual (left) and V1 (right) to prefrontal cortices in sighted adults, blind adults, and sighted infants. Blind adults show a larger difference than any of the other groups. Asterisks (*) denote significant Bonferroni-corrected pairwise comparisons (p < .05).

Revised version: Bar graph shows within hemisphere (blue) and between hemisphere (orange) functional connectivity (r coefficient of resting state correlations) of secondary visual (left) and V1 (right) to prefrontal cortices in sighted adults, blind adults, and infants. Blind adults show a larger difference than any of the other groups. Asterisks (*) denote significant Bonferroni-corrected pairwise comparisons (p < .05).

Task-based fMRI studies find that cross-modal responses in occipital cortex colateralize with fronto-parietal networks with related functions (e.g., language, response selection) (Kanjlia et al., 2021; Lane et al., 2017). For example, language-responsive occipital areas collateralize with language responsive prefrontal areas across individuals (Lane et al., 2017). Recruitment of visual cortices by cross-modal tasks (e.g., spoken language) may enhance within-hemisphere connectivity in people born blind (Kanjlia et al., 2021; Lane et al., 2017; Tian et al., 2022). Together, this evidence supports the hypothesis that, starting from the less lateralized infant state, blindness increases lateralization of occipital long-range connectivity.

Specialization of connectivity across different fronto-occipital networks: present in adults, absent at birth

In blind adults, different occipital areas show enhanced connectivity patterns with distinct subregions of PFC and this specialization is aligned with the functional specialization observed in task-based data (Bedny et al., 2011; Kanjlia et al., 2016, 2021). For example, language-responsive subregions of occipital cortex show strongest functional connectivity with language-responsive sub-regions of PFC, whereas math-responsive occipital areas show stronger connectivity with math-responsive PFC. This pattern is most pronounced in blind people but can be seen weakly even in sighted participants (Fig. 4) (Bedny et al., 2011; Kanjlia et al., 2016, 2021; Lane et al., 2015). Is this fronto-occipital connectivity specialization present in infancy, potentially enabling the task-based cross-modal specialization?

Occipito-frontal functional connectivity.

Earlier version: Bar graph shows across functional connectivity of different sub-regions of prefrontal (PFC) and occipital cortex (OCC) in sighted adults, blind adults, and sighted infants. Sub-regions (regions of interest) were defined based on task-based responses in a separate dataset of sighted (frontal) and blind (frontal and occipital) adults (Kanjlia et al., 2016, 2021; Lane et al., 2015). PFC/OCC-MATH: math-responsive regions were more active when solving math equations than comprehending sentences. PFC/OC-LANG: language-responsive regions were more active when comprehending sentences than solving math equations (Kanjlia et al., 2016, 2021; Lane et al., 2015). In blind adults these regions show biases in connectivity related to their function i.e., language-responsive PFC is more correlated with language responsive OCC. No such pattern is observed in infants. Asterisks (*) denote significant Bonferroni-corrected pairwise comparisons (p < .05). See Supplementary Figure S9 for connectivity matrix.

Revised version: Bar graph shows across functional connectivity of different sub-regions of prefrontal (PFC) and occipital cortex (OCC) in sighted adults, blind adults, and infants. Sub-regions (regions of interest) were defined based on task-based responses in a separate dataset of sighted (frontal) and blind (frontal and occipital) adults (Kanjlia et al., 2016, 2021; Lane et al., 2015). PFC/OCC-MATH: math-responsive regions were more active when solving math equations than comprehending sentences. PFC/OC-LANG: language-responsive regions were more active when comprehending sentences than solving math equations (Kanjlia et al., 2016, 2021; Lane et al., 2015). In blind adults these regions show biases in connectivity related to their function i.e., language-responsive PFC is more correlated with language responsive OCC. No such pattern is observed in infants. Asterisks (*) denote significant Bonferroni-corrected pairwise comparisons (p < .05). See Supplementary Figure S9 for connectivity matrix.

We compared connectivity preferences across three prefrontal and three occipital regions previous shown to activate preferentially in language (sentences > math), math (math > sentences) and response-conflict (no-go > go with tones) (Kanjlia et al., 2016, 2021; Lane et al., 2015). For ease of viewing, Fig. 4 shows results from two of the three regions, math and language. See Supplementary Figure S8 for all three regions. Note that the statistical analyses included all three areas.

Earlier version: Contrary to the hypothesis that specialization of functional connectivity across different prefrontal/occipital areas is present from birth, sighted infants showed a less differentiated fronto-occipital connectivity pattern relative to both blind and sighted adults (Group (sighted adults, blind adults, sighted infants) by occipital regions (math, language, response-conflict) by PFC regions (math, language, response-conflict) interaction F(8, 2208) = 16.323, p < 0.001). Unlike in adults, in infants, all the occipital regions showed stronger correlations with math- and response-conflict related prefrontal areas than language-responsive prefrontal areas (Fig. 4 and Supplementary Figure S8). However, the preferential correlation with math responsive PFC was strongest in those occipital areas that go on to develop math responses in blind adults (occipital regions (math, language, response-conflict) by PFC regions (math, language, response-conflict) interaction in infants F(4, 1896) = 85.145, p < 0.001, post-hoc Bonferroni-corrected paired t-test see Supplementary Table 1).

Revised version: Contrary to the hypothesis that specialization of functional connectivity across different prefrontal/occipital areas is present from birth, infants showed a less differentiated fronto-occipital connectivity pattern relative to both blind and sighted adults (Group (sighted adults, blind adults, infants) by occipital regions (math, language, response-conflict) by PFC regions (math, language, response-conflict) interaction F(8, 2208) = 16.323, p < 0.001). Unlike in adults, in infants, all the occipital regions showed stronger correlations with math- and response-conflict related prefrontal areas than language-responsive prefrontal areas (Fig. 4 and Supplementary Figure S8). However, the preferential correlation with math responsive PFC was strongest in those occipital areas that go on to develop math responses in blind adults (occipital regions (math, language, response-conflict) by PFC regions (math, language, response-conflict) interaction in infants F(4, 1896) = 85.145, p < 0.001, post-hoc Bonferroni-corrected paired t-test see Supplementary Table S1).

Note that findings of reduced regional specialization in infants need to be interpreted with caution. First, we do not know whether the same specialization of prefrontal sub-regions seen in adults is present in infants, although, prior evidence suggests some prefrontal specialization is already present (Raz & Saxe, 2020). Second, the more fine-grained comparisons across occipital/frontal regions are more vulnerable to potential anatomical alignment issues between adult and infant brains. In other words, lack of specialization in infants could reflect the different location of the areas in this population.

Discussion

Earlier version: The present results provide insight into the developmental process of experience-based functional specialization in human cortex. We find independent effects of visual experience and blindness on the development of visual networks. Aspects of the sighted adult connectivity pattern require visual experience to become established. This is particularly striking for secondary visual cortices, where, connectivity with nonvisual networks in sighted infants resemble blind more than sighted adults. Both in infants and blind adults, secondary occipital areas showed stronger functional connectivity with higher-order prefrontal cortices than with other sensory-motor networks (S1/M1, A1). Consistent with this observation, one previous study with a small sample of infants found strong connectivity between lateral occipital and prefrontal areas, although there was no comparison to blind adults in that study (Barttfeld et al., 2018). In V1, infants fell somewhere in between sighted and blind adults, suggesting an effect both of vision and of blindness on functional connectivity.

Revised version: The present results provide insight into the developmental process of experience-based functional specialization in human cortex. We find independent effects of visual experience and blindness on the development of visual networks. Aspects of the sighted adult connectivity pattern require visual experience. This is particularly striking for secondary visual cortices, where, connectivity with nonvisual networks in infants resemble blind more than sighted adults. Both in infants and blind adults, secondary occipital areas showed stronger functional connectivity with higher-order prefrontal cortices than with other sensory-motor networks (S1/M1, A1). Consistent with this observation, one previous study with a small sample of infants found strong connectivity between lateral occipital and prefrontal areas, although there was no comparison to blind adults in that study (Barttfeld et al., 2018). In V1, infants fell somewhere in between sighted and blind adults, suggesting an effect both of vision and of blindness on functional connectivity.

The present results reveal the effects of experience on development of functional connectivity between infancy and adulthood, but do not speak to the precise time course of these effects. Infants in the current sample had between 0 and 20 weeks of visual experience. Comparisons across these infants suggests that several weeks of postnatal visual experience is insufficient to produce a sighted-adult connectivity profile. The time course of development could be anywhere between a few months and years and could be tested by examining data from children of different ages.

Added in the revised version: We propose that vision, as well as temporally coordinated multi-modal experiences contribute to establishing the sighted adult connectivity profile in visual cortices. For example, coordinated visuo-motor activity during development may enhance connectivity between visual and motor networks. Supporting this conjecture, in infants, motor competence and early experience predict coupling between occipital and motor networks (Colomer et al., 2023).

Earlier version: We propose that vision, as well as temporally coordinated multi-modal experiences contribute to establishing the sighted adult connectivity profile in visual cortices. For example, coordinated visuo-motor activity during development may enhance connectivity between visual and motor networks. Supporting this idea, evidence shows that in infants, motor competence and early experience predict coupling between occipital and motor networks (Colomer et al., 2023).Despite these insights, many questions remain regarding the neurobiological mechanisms underlying experience-based functional connectivity changes and their relationship to anatomical development. Long-range anatomical connections between brain regions are already present in infants—even prenatally—though they remain immature (Huang et al., 2009; Kostović et al., 2019, 2021; Takahashi et al., 2012; Vasung, 2017). Functional connectivity changes may stem from local synaptic modifications within these stable structural pathways, consistent with findings that functional connectivity can vary independently of structural connection strength (Fotiadis et al., 2024). Moreover, functional connectivity has been shown to outperform structural connectivity in predicting individual behavioral differences, suggesting that experience-based functional changes may reflect finer-scale synaptic or network-level modulations not captured by macrostructural measures (Ooi et al., 2022). Prior studies also suggest that, even in adults, coordinated sensory-motor experience can lead to enhancement of functional connectivity across sensory-motor systems, indicating that large-scale changes in functional connectivity do not necessarily require corresponding changes in anatomical connectivity (Guerra-Carrillo et al., 2014; Li et al., 2018). Another possibility is some connectivity changes are mediated by ‘third-party’ regions, including subcortical mechanisms such as those involving the thalamus (Vega-Zuniga et al., 2025).

Revised version: Many questions remain regarding the neurobiological mechanisms underlying experience-based functional connectivity changes and their relationship to anatomical development. Resting state functional correlations are an indirect measure of both function and anatomy and differences in these measures are consistent with many possible underlying biological mechanisms. Long-range anatomical connections between brain regions are already present in infants—even prenatally—though they remain immature (Huang et al., 2009; Kostović et al., 2019, 2021; Takahashi et al., 2012; Vasung, 2017). Functional connectivity changes may stem from local synaptic modifications within these stable structural pathways, consistent with findings that functional connectivity can vary independently of structural connection strength (Fotiadis et al., 2024). Moreover, functional connectivity has been shown to outperform structural connectivity in predicting individual behavioral differences, suggesting that experience-based functional changes may reflect finer-scale synaptic or network-level modulations not captured by macrostructural measures (Ooi et al., 2022). Prior studies also suggest that, even in adults, coordinated sensory-motor experience can lead to enhancement of functional connectivity across sensory-motor systems, indicating that large-scale changes in functional connectivity do not necessarily require corresponding changes in anatomical connectivity (Guerra-Carrillo et al., 2014; Li et al., 2018). Resting-state functional connectivity captures synchrony in BOLD signal fluctuations rather than causal interactions and differences in functional connectivity cannot on their own reveal how underlying neurophysiological mechanisms are modified. Connectivity changes between two areas could be mediated by ‘third-party’ hub regions. For example, posterior parietal cortex serves as a cortical hub for multisensory integration and visuo-motor coordination and could mediate occipital-to-sensory-motor communication (Rolls et al., 2023; Sereno & Huang, 2014). Subcortical structures such as the thalamus could also play a mediating role (Vega-Zuniga et al., 2025). Future studies will be needed to determine whether these functional changes are accompanied by alterations in structural connectivity, and to probe causal interactions and mechanistic underpinnings.

Earlier version: The current findings reveal both effects of vision and effects of blindness on the functional connectivity patterns of the visual cortex. A further open question is whether visual experience plays an instructive or permissive role in shaping neural connectivity patterns. An instructive role suggests that specific sensory experiences or patterns of neural activity directly shape and organize neural circuitry. In contrast, a permissive role implies that sensory experience or neural activity merely facilitates the influence of other factors—such as molecular signals—on the formation and organization of neural circuits (Crair, 1999; Sur et al., 1999). Studies with animals that manipulate the pattern or informational content of neural activity while keeping overall activity levels constant could distinguish between these hypotheses (Crair, 1999; Roy et al., 2020; Stellwagen & Shatz, 2002).

Revised version: The current findings reveal both effects of vision and effects of blindness on the functional connectivity patterns of the visual cortex. A further open question is whether visual experience plays an instructive or permissive role in shaping neural connectivity patterns. An instructive role implies that sensory experiences or patterns of neural activity directly shape and organize neural circuitry. In contrast, a permissive role implies that sensory experience or neural activity merely facilitates the influence of other factors—such as molecular signals—on the formation and organization of neural circuits (Crair, 1999; Sur et al., 1999). Studies with animals that manipulate the pattern or informational content of neural activity while keeping overall activity levels constant could distinguish between these hypotheses (Crair, 1999; Roy et al., 2020; Stellwagen & Shatz, 2002). In humans, such manipulations are not feasible, leaving us to study only the consequences of the presence or absence of vision. Under an instructive account, visual and multisensory experience could strengthen coupling between visual and other non-visual sensory-motor cortices through coordinated activity, thereby establishing the sighted-adult connectivity pattern. In the absence of visual input, by contrast, the lack of such coordinated activity may prevent these couplings from being established. Alternatively, vision may act permissively, indirectly enabling maturational processes that shift connectivity toward the sighted-adult configuration.

Earlier version: A further key question concerns the behavioral relevance of the connectivity signatures observed in the current study. The capacity of occipital cortices to support visual and multimodal behavior in sighted people may depend not only on local visual cortex function but also on the capacity of the visual system to coordinate its function with non-visual networks. Does enhanced connectivity between visual and non-visual sensory motor networks facilitate multimodal integration for sighted people e.g., when catching a ball? Potentially consistent with this possibility, recent evidence suggests that people who grew up blind but recover sight in adulthood show multimodal integration deficits (Ashtari, 2020; Badde et al., 2020; Guerreiro et al., 2015; Putzar et al., 2007) and distinct occipital oscillations (Pant et al., 2023).

Revised version: A further key question concerns the behavioral relevance of the connectivity signatures observed in the current study. The capacity of occipital cortices to support visual and multimodal behavior in sighted people may depend not only on local visual cortex function but also on the capacity of the visual system to coordinate its function with non-visual networks. Does enhanced connectivity between visual and non-visual sensory motor networks facilitate multimodal integration for sighted people e.g., when catching a ball? Potentially consistent with this possibility, recent evidence suggests that people who grew up blind but recover sight in adulthood show multimodal integration deficits (Badde et al., 2020; Guerreiro et al., 2015; Mowad et al., 2020; Putzar et al., 2007) and distinct occipital oscillations (Pant et al., 2023).

Earlier version: Conversely, for people who remain blind throughout life, visual-PFC connectivity could enable recruitment of visual cortices for higher-order non-visual functions, such as language and executive control. Habitual activation of occipital networks during higher cognitive tasks in early development could intern enhance connectivity and specialization of visual networks for non-visual functions.

Revised version: Conversely, for people who remain blind throughout life, visual-PFC connectivity could enable recruitment of visual cortices for higher-order non-visual functions, such as language and executive control (Bedny et al., 2011; Kanjlia et al., 2021). Our results suggest that the pattern of connectivity observed in blind adults may build on connectivity patterns already present in infancy: like blind adults, infants show stronger occipital–PFC than occipital–sensory–motor coupling. Repeated engagement of occipital networks during higher cognitive tasks in early development could further enhance connectivity and enable specialization of visual networks for different non-visual higher-order functions.

Added in the revised version: Some prior studies have measured resting-state and task-based functional profiles in the same participants. These studies find that within visual cortices of blind people, the task-based profile of a cortical area is related to its resting state connectivity pattern (Abboud & Cohen, 2019b; Deen et al., 2015; Kanjlia et al., 2016, 2021). This suggests that these two measures are related. However, the timecourse of this relationship, the developmental trajectory and mechanism of plasticity is not known. Primarily this is because there is very little relevant developmental evidence. For example, in the current study we find that the resting state profile of secondary visual networks in infants is similar to that of blind adults. However, we do not know whether the visual cortices of infants show enhanced task-based cross modal responses, relative to sighted adults and how this compares to responses observed in blind adults. Future work with infants and children would be able to address this question.

Earlier version: In the current study, the clearest evidence for functional change driven by blindness was observed for laterality. Connectivity lateralization in sighted infants resembles that of sighted adults, in both V1 and secondary visual cortices. Relative to both sighted infants and sighted adults, blind adults show more lateralized connectivity patterns between occipital and prefrontal cortices. Previous studies suggest that in people born blind occipital and non-occipital language responses are co-lateralized (Lane et al., 2017). We speculate that habitual activation of visual cortices by higher-cognitive tasks, such as language, which are themselves highly lateralized, contributes to this biased connectivity pattern of occipital cortex in blindness.

Revised version: In the current study, the clearest evidence for functional change driven by blindness was observed for laterality. Connectivity lateralization in sighted infants resembles that of sighted adults, in both V1 and secondary visual cortices. Relative to both sighted infants and sighted adults, blind adults show more lateralized connectivity patterns between occipital and prefrontal cortices. Previous studies suggest that in people born blind occipital and non-occipital language responses are co-lateralized (Lane et al., 2017; Tian et al., 2023). We speculate that habitual activation of visual cortices by higher-cognitive tasks, such as language, which are themselves highly lateralized, contributes to this biased connectivity pattern of occipital cortex in blindness. Taken together, these results suggest a developmental framework in which intrinsic connectivity present in infancy provides a scaffold that is subsequently shaped and reinforced by experience-dependent recruitment, through either visual experience or the lifelong absence of vision in blindness. Longitudinal work across successive developmental stages will be crucial to test how the alternative trajectories shaped by visual experience versus blindness unfold over development.

Materials and methods

Participants

Fifty sighted adults and thirty congenitally blind adults contributed the resting state data (sighted: n = 50; 30 females; mean age = 35.33 years, standard deviation (SD) = 14.65; mean years of education = 17.08, SD = 3.1; blind: n = 30; 19 females; mean age = 44.23 years, SD = 16.41; mean years of education = 17.08, SD = 2.11; blind vs. sighted age, t (78) = 2.512, p < 0.05; blind vs. sighted years of education, t (78) = 0.05, p = 0.996). Since blind participants were on average older, we also performed analyses in an age-matched subgroups of sighted controls (n = 29) and found similar results to the full sample (see Supplementary Figure S10). Blind and sighted participants had no known cognitive or neurological disabilities (screened through self-report). All adult anatomical images were read by a board-certified radiologist, and no gross neurological abnormalities were found. All the blind participants had at most minimal light perception from birth. Blindness was caused by pathology anterior to the optic chiasm (i.e., not due to brain damage). All participants gave written informed consent under a protocol approved by the Institutional Review Board of Johns Hopkins University.

Earlier version: Neonate data were from the third release of the Developing Human Connectome Project (dHCP) (n = 783) (https://www.developingconnectome.org). Ethical approval was obtained from the UK Health Research Authority (Research Ethics Committee reference number: 14/LO/1169). After quality control procedures (described below), 475 subjects were included in data analysis, with one scan per subject. The average age from birth at scan = 2.79 weeks (SD = 3.77, median = 1.57, range = 0 – 19.71); average gestational age at scan = 41.23 weeks (SD = 1.77, median = 41.29, range = 37 – 45.14); average gestational age at birth = 38.43 weeks (SD = 3.73, median = 39.71, range = 23 – 42.71). We only included infants who were full-term or scanned at term-equivalent age if preterm, while not being flagged by the dHCP project team as not passing quality control for functional MRI (fMRI) images (n = 634). Infants with more than 160 motion outliers were exclude (n = 116 dropped). Motion-outlier volumes were defined as DVARS (the root mean square intensity difference between successive volumes) higher than 1.5 interquartile range above the 75th centile, after motion and distortion correction. Infants with signal drop-out in regions of interest (ROI) were also excluded (n = 43 dropped). To identify signal dropout, we first averaged blood oxygen level-dependent (BOLD) signal intensity for all time point, for each subject, in each of 100 parcel defined by Schaefer’s atlas (Schaefer et al., 2018). For each ROI (n = 18 ROIs) in the current study, signal dropout was then identified as BOLD intensity lower than -3 standard deviations, where the mean and standard deviations were identified across all 100 cortical parcels. Participants were excluded if any of the ROIs showed a signal dropout. The infants’ structural images were reviewed by a pediatric neuroradiologist from the dHCP team, who assigned scores on a scale from 1 to 5. A score of 1 indicated a normal appearance for the subject’s age, while scores of 4 or 5 suggested potential or likely clinical significance, or both clinical and imaging relevance. We repeated our analysis after excluding infants with a radiology score of 4 or 5, and the results remained consistent (see Figure S11).

Revised version: Neonate data were from the third release of the Developing Human Connectome Project (dHCP) (n = 783) (https://www.developingconnectome.org). Ethical approval was obtained from the UK Health Research Authority (Research Ethics Committee reference number: 14/LO/1169). After quality control procedures (described below), 475 subjects were included in data analysis, with one scan per subject. The average age from birth at scan = 2.79 weeks (SD = 3.77, median = 1.57, range = 0 – 19.71); average gestational age at scan = 41.23 weeks (SD = 1.77, median = 41.29, range = 37 – 45.14); average gestational age at birth = 38.43 weeks (SD = 3.73, median = 39.71, range = 23 – 42.71). We only included infants who were full-term or scanned at term-equivalent age if preterm, while not being flagged by the dHCP project team as not passing quality control for functional MRI (fMRI) images (n = 634). Infants with more than 160 motion outliers were exclude (n = 116 dropped). Motion-outlier volumes were defined as DVARS (the root mean square intensity difference between successive volumes) higher than 1.5 interquartile range above the 75th centile, after motion and distortion correction. Infants with signal drop-out in regions of interest (ROI) were also excluded (n = 43 dropped). To identify signal dropout, we first averaged blood oxygen level-dependent (BOLD) signal intensity for all time point, for each subject, in each of 100 parcel defined by Schaefer’s atlas (Schaefer et al., 2018). For each ROI (n = 18 ROIs) in the current study, signal dropout was then identified as BOLD intensity lower than -3 standard deviations, where the mean and standard deviations were identified across all 100 cortical parcels. Participants were excluded if any of the ROIs showed a signal dropout. The same signal dropout assessment was also applied to the blind and sighted adults to ensure consistent quality control across groups. One participant in the sighted adult group and two participants in the blind adult group exhibited signal dropout in one ROI each. Excluding these participants did not alter the group-level results (see Figure S13). The infants’ structural images were reviewed by a pediatric neuroradiologist from the dHCP team, who assigned scores on a scale from 1 to 5. A score of 1 indicated a normal appearance for the subject’s age, while scores of 4 or 5 suggested potential or likely clinical significance, or both clinical and imaging relevance. We repeated our analysis after excluding infants with a radiology score of 4 or 5, and the results remained consistent (see Figure S11).

Image acquisition

Blind and sighted adult MRI anatomical and functional images were collected on a 3T Phillips scanner at the F. M. Kirby Research Center. T1-weighted anatomical images were collected using a magnetization-prepared rapid gradient-echo (MP-RAGE) in 150 axial slices with 1 mm isotropic voxels. Resting state fMRI data were collected in 36 sequential ascending axial slices for 8 minutes. TR = 2 s, TE = 0.03 s, flip angle = 70°, voxel size = 2.4 × 2.4 × 2.5 mm, inter-slice gap = 0.5 mm, field of view (FOV) = 192 × 172.8 × 107.5. Participants completed 1 to 4 scans of 240 volume each (average scan time = 710.4 second per person). During the resting state scan, participants were instructed to relax but remain awake. Sighted participants wore light-excluding blindfolds to equalize the light conditions across the groups during the scans.

Infants (dHCP) Anatomical and functional images were collected on a 3T Phillips scanner at the Evelina Newborn Imaging Centre, St Thomas’ Hospital, London, UK. A dedicated neonatal imaging 219 system including a neonatal 32-channel phased-array head coil was used. T2w multi-slice fast spin-echo images were acquired with in-plane resolution 0.8×0.8 mm2 and 1.6 mm slices overlapped by 0.8 mm (TR = 12000 ms, TE = 156 ms, SENSE factor 2.11 axial and 2.6 sagittal). In infants, T2w images were used as the anatomical image because the brain anatomy is more clearly in T2w than in T1w images. Fifteen minutes of resting state fMRI data were collected using a used multiband 9x accelerated echo-planar imaging (TR = 392 ms, TE = 38 ms, 2300 volumes, with an acquired resolution of 2.15 mm isotropic). Single-band reference scans were acquired with bandwidth-matched readout, along with additional spin-echo acquisitions with both AP/PA fold-over encoding directions.

Data analysis

Resting state data were preprocessed using FSL version 5.0.9 (Smith et al., 2004), DPABI version 6.1 (Yan et al., 2016), FreeSurfer (Dale et al., 1999), and in-house code (https://github.com/NPDL/Resting-state_dHCP). The functional data for all groups were linearly detrended and low-pass filtered (0.08 Hz).

For adults, functional images were registered to the T1-weighted structural images, motion corrected using MCFLIRT (Jenkinson et al., 2002), and temporally high-pass filtering (150 s). No subject had excessive head movement (> 2mm) or rotation (> 2°) at any timepoint. Resting state data are known to include artifacts related to physiological fluctuations such as cardiac pulsations and respiratory-induced modulation of the main magnetic field. A component-based method, CompCor (Behzadi et al., 2007), was therefore used to control for these artifacts. Particularly, following the procedure described in Whitfield-Gabrieli et al., nuisance signals were extracted from 2-voxel eroded masks of spinal fluid (CSF) and white matter (WM), and the first 5 principal components analysis (PCA) components derived from these signals was regressed out from the processed BOLD time series (Whitfield-Gabrieli & Nieto-Castanon, 2012). In addition, a scrubbing procedure was applied to further reduce the effect of motion on functional connectivity measures (Power et al., 2012, 2014). Frames with root mean square intensity difference exceeding 1.5 interquartile range above the 75th centile, after motion and distortion correction, were censored as outliers.

Earlier version: The infants resting state functional data were pre-processed by the dHCP group using the project’s in-house pipeline (Fitzgibbon et al., 2020), This pipeline uses a spatial independent component analysis (ICA) denoising step to minimize artifact due to multi-band artefact, residual head-movement, arteries, sagittal sinus, CSF pulsation. For infants, ICA denoising is preferable to using CSF/WM regressors. Because it is challenging to accurately define anatomical boundaries of CSF/WM due to the low imaging resolution comparing with the brain size and the severe partial-volume effect in the neonate (Fitzgibbon et al., 2020). Like in the adults, frames with root mean square intensity difference exceeding 1.5 interquartile range above the 75th centile, after motion and distortion correction, were considered as motion outliers. Out from the 2300 frames, a subset of continuous 1600 with minimum number of motion outliers was kept for each subject. Motion outliers were censored from the subset of continuous 1600, and a subject was excluded from further analyses when the number of outlier exceeded 160 (10% of the continues subset) (Hu et al., 2022). While infant connectivity estimates may be less robust at the individual level compared to adults due to shorter scan durations and higher motion, our cohort’s large sample size (n=475) and rigorous motion censoring mitigate these limitations for group-level analyses. Substantial differences between the groups exist in this study, including the number of subjects, brain sizes, imaging parameters, and data preprocessing, all of which are likely to have an impact on the overall signal quality. To address this concern, we compared the split-half noise ceiling across the groups (infants, sighted adults, and blind adults). For each participant, we first split the rs-fMRI time series into two halves, then calculated the ROI-wise rsFC pattern from the two splits. The split-half noise ceiling was estimated according to Lage-Castellanos et al (Lage-Castellanos et al., 2019). The noise ceilings of the three groups (infants: 0.90 ± 0.056, blind adults: 0.88 ± 0.041, sighted adults: 0.90 ± 0.055) showed no significant difference (One-way ANOVA, F (2,552) = 2.348, p = 0.097). Therefore, overall signal quality is unlikely to impact our results.

Revised version: The infants resting state functional data were pre-processed by the dHCP group using the project’s in-house pipeline (Fitzgibbon et al., 2020), This pipeline uses a spatial independent component analysis (ICA) denoising step to minimize artifact due to multi-band artefact, residual head-movement, arteries, sagittal sinus, CSF pulsation. For infants, ICA denoising is preferable to using CSF/WM regressors. Because it is challenging to accurately define anatomical boundaries of CSF/WM due to the low imaging resolution comparing with the brain size and the severe partial-volume effect in the neonate (Fitzgibbon et al., 2020). Like in the adults, frames with root mean square intensity difference exceeding 1.5 interquartile range above the 75th centile, after motion and distortion correction, were considered as motion outliers. Out from the 2300 frames, a subset of continuous 1600 with minimum number of motion outliers was kept for each subject. Motion outliers were censored from the subset of continuous 1600, and a subject was excluded from further analyses when the number of outlier exceeded 160 (10% of the continues subset) (Hu et al., 2022). While infant connectivity estimates may be less robust at the individual level compared to adults due to shorter scan durations and higher motion, our cohort’s large sample size (n=475) and rigorous motion censoring mitigate these limitations for group-level analyses. Substantial differences between the groups exist in this study, including the number of subjects, brain sizes, imaging parameters, and data preprocessing, all of which are likely to have an impact on the overall signal quality. To assess the reliability of the ROI-wise FC used in all analyses, we computed split-half noise ceilings for each participant. The rs-fMRI time series was divided into two equal halves, ROI-wise FC profiles were calculated separately for each half, and the noise ceiling for each ROI was defined as the Pearson correlation between the two profiles following Lage-Castellanos et al. (2019). We then averaged the ROI-wise noise ceilings for each group. The resulting values (infants: 0.90 ± 0.056; blind adults: 0.88 ± 0.041; sighted adults: 0.90 ± 0.055) did not differ significantly (one-way ANOVA: F(2,552) = 2.348, p = 0.097). We also examined noise ceilings for each ROI separately. All ROIs showed high absolute reliability (noise ceiling > 0.80) across groups. Although many ROIs showed statistically significant group differences in noise ceiling (one-way ANOVA, p < 0.05), the effect sizes were small to moderate (partial η² < 0.14). These findings indicate that reliability may vary modestly across groups at the ROI level. We cannot determine whether such variability contributes to the group differences reported in this study, but the consistently high absolute reliability suggests that the FC estimates are generally robust. The full ROI-wise statistical results are provided in the Supplementary Table S2. For both groups of adult and infants, we performed a temporal low-pass filter (0.08 Hz low-pass cutoff) and a linear detrending. ROI-to-ROI connectivity was calculated using Pearson’s correlation between ROI-averaged BOLD timeseries (ROI definition see below). The All t-tests and F-tests are two-sided. The comparison of correlation coefficients was done using cocor software package and Pearson and Filon’s z (Diedenhofen & Musch, 2015; Pearson & Filon, 1898).

Removed in the revised version: For both groups of adult and infants, we performed a temporal low-pass filter (0.08 Hz low-pass cutoff) and a linear detrending. ROI-to-ROI connectivity was calculated using Pearson’s correlation between ROI-averaged BOLD timeseries (ROI definition see below). The All t-tests and F-tests are two-sided. The comparison of correlation coefficients was done using cocor software package and Pearson and Filon’s z (Diedenhofen & Musch, 2015; Pearson & Filon, 1898).

ROI definition

Frontal and secondary visual ROIs were defined functionally based on data from a separate task-based fMRI experiments with blind and sighted adults (Kanjlia et al., 2016, 2021; Lane et al., 2015). Three separate experiments were conducted with the same group of blind and sighted subjects (sighted n=18; blind n=23). The language ROIs in the occipital and frontal cortices were identified by sentence > nonwords contrast in an auditory language comprehension task (Lane et al., 2015). The math ROIs were identified by math > sentence contrast in an auditory task where participants judged equivalence of pairs of math equations and pairs of sentences (Kanjlia et al., 2016). The executive function ROIs were identified by no-go > frequent go contrast in an auditory go/no-go task with non-verbal sounds (Kanjlia et al., 2021). All ROI files are available at openICPSR (https://doi.org/10.3886/E198832V1).

Earlier version: Occipital secondary ‘visual’ ROIs were defined based on group comparisons blind > sighted in a whole-cortex analysis (Figure 1, Figure S5.) The occipital language ROI was defined as the cluster that responded more to auditory sentence than auditory nonwords conditions in blind, relative to sighted, in a whole-cortex analysis, likewise the occipital math ROI was defined as math>sentences, blind>sighted interaction and the occipital executive ROI as no-go > frequent go, blind>sighted. All three occipital ROIs were defined in the right hemisphere. Left-hemisphere occipital ROIs were created by flipping the right-hemisphere ROIs to the left hemisphere. Each functional ROI spans multiple anatomical regions and together the secondary visual ROIs tile large portions of lateral occipital, occipito-temporal, dorsal occipital and occipito-parietal cortices. In sighted people, the secondary visual occipital ROIs include the anatomical locations of functional regions such as motion area V5/MT+, the lateral occipital complex (LO), category specific ventral occipitotemporal cortices and dorsally, V3a and V4v. The occipital ROI also covers the middle of the ventral temporal lobe. Dorsally, it extended to the intraparietal sulcus and superior parietal lobule.

Revised version: Occipital secondary ‘visual’ ROIs were defined based on group comparisons blind > sighted in a whole-cortex analysis (Figure 1, Figure S5.) The occipital language ROI was defined as the cluster that responded more to auditory sentence than auditory nonwords conditions in blind, relative to sighted, in a whole-cortex analysis, likewise the occipital math ROI was defined as math>sentences, blind>sighted interaction and the occipital executive ROI as no-go > frequent go, blind>sighted. All three occipital ROIs were defined in the right hemisphere. Left-hemisphere occipital ROIs were created by flipping the right-hemisphere ROIs to the left hemisphere. Each functional ROI spans multiple anatomical regions and together the secondary visual ROIs tile large portions of lateral occipital, occipito-temporal, dorsal occipital and occipito-parietal cortices. In sighted people, the secondary visual occipital ROIs include the anatomical locations of functional regions such as motion area V5/MT+, the lateral occipital complex (LO), as well as ventral portions encompassing category-selective ventral occipito-temporal cortices, including V4v, and dorsal portions including V3a. The occipital ROI also extended ventrally into the middle portion of the temporal lobe and dorsally into the intraparietal sulcus and superior parietal lobule. The frontal PFC ROIs were defined functionally, based on a whole-cortex analysis which combined all blind and sighted adult data. The frontal language ROI was defined as responded more auditory sentence than auditory nonwords conditions across all blind and sighted subjects, constrained to the prefrontal cortex. Likewise, math responsive PFC was defined as math>sentences and executive no-go > frequent go. For frontal ROIs, the language ROI was defined in the left, and the math and executive function ROI were defined in the right hemisphere, then flip to the other hemisphere.

Removed in the revised version: The frontal PFC ROIs were defined functionally, based on a whole-cortex analysis which combined all blind and sighted adult data. The frontal language ROI was defined as responded more auditory sentence than auditory nonwords conditions across all blind and sighted subjects, constrained to the prefrontal cortex. Likewise, math responsive PFC was defined as math>sentences and executive no-go > frequent go. For frontal ROIs, the language ROI was defined in the left, and the math and executive function ROI were defined in the right hemisphere, then flip to the other hemisphere.

The V1 ROI was defined from a previously published anatomical surface-based atlas (PALS-B12) (Van Essen, 2005). The primary somatosensory and motor cortex (S1/M1) ROI was selected as the area that responds more to the go than no-go trials in the auditory go/no-go task across both blind and sighted groups, constrained to the hand area in S1/M1 search-space from neurosynth.org (term “hand movements”) (Kanjlia et al., 2021). The primary auditory cortex (A1) ROI was defined as the transverse temporal portion of a gyral-based atlas (Desikan et al., 2006; Morosan et al., 2001).

All the ROIs were defined in standard space and then transformed into each subject’s native space. For adults this was done by employing the deformation field estimated by FreeSurfer. For infants, the ROIs were transformed into each subject’s native space using a two-step approach. First, the ROIs were converted from the adult’s MNI space into the 40-week dHCP template (Bozek et al., 2018). ANTS, previously shown to be effective in pediatric studies (Avants et al., 2014; Cabral et al., 2022; Jain et al., 2012; Lawson et al., 2013), was utilized to estimate the deformation field between these two spaces. In this step, the infant’s scalp and cerebellum were masked, as these structures in the infant brain greatly differ from those in the adult and can introduce bias into the registration process, as outlined in a study by Cabral et al (Cabral et al., 2022). Secondly, the ROIs were further transformed from the 40-week template space into each individual’s native spaces, employing the deformation field provided by the dHCP group. Nearest neighbor interpolation was applied in both steps (Examples of the resulting ROI alignment on individual brains are shown in Supplementary Figure S12). For both adults and infants, any overlapping voxels between ROIs were removed and not counted toward any ROIs.